---
title: 'Creating an output'
description: 'How to use the print function to create an output using Python'
---

## Example 1

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

The example below demonstrates how to code a very simple output using Python

`@instructions`
All you have to do here is take a look at the code and then click on RUN

`@hint`
You just click on the run button!

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
print("Hi")
```

`@solution`
```{python}
print("Hi")
```

`@sct`
```{python}
Ex().has_output(r'Hi')
success_msg("Well Done!")
```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: ffc6e00c2d
xp: 100
```

Outputting text using the print function.

`@instructions`
Python uses the `print()` function to output information to the screen.
Anything between the parenthesis will appear on the screen.
Try to print out your name in the windows below.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
