---
title: 'Creating an output'
description: 'How to use the print function to create an output using Python'
---

## Example 1

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

The example below demonstrates how to code a very simple output using Python

`@instructions`
All you have to do here is take a look at the code and then click on RUN

`@hint`
You just click on the run button!

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
print("Hi")
```

`@solution`
```{python}
print("Hi")
```

`@sct`
```{python}
Ex().has_printout("Hi")
msg = "Well done"
```
